---
title: blogdown 패키지 설치 및 환경설정
description: >-
  Since my site is already created (and I didn’t think about documenting it while I was doing it), this post documents the creation of a near-identical site for an animation studio: Animoplex.
I’m erring on the side of too much detail in this post, so skip past things you know as needed.
I know, I’m behind the times and am still using the command line outside of RStudio to work with GitHub. Some day I’ll convert. If you’d like to do that today, there’s information here.


---
## 연습문제 1.1

```yaml
type: NormalExercise
lang: r
xp: 100
skills: 1
key: 9ed0377aec
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}
devtools::install_github("rstudio/blogdown")
library(blogdown)
```
`@sample_code`
```{r}
devtools::install_github("rstudio/blogdown")
library(blogdown)
```

`@sct`
```{r}
devtools::install_github("rstudio/blogdown")
library(blogdown)
```



